# Assignment 1.4 Pandas
Pandas is an open source Python package providing fast, flexible, and expressive data structures and data manipulation designed to make working with “relational” or “labeled” data both easy and intuitive. It is the fundamental high-level building block for doing practical, real world data analysis in Python.
https://pandas.pydata.org/

<!-- grade_id: cell-462d89bbf25b800a, points: 0, grade: false, solution: true -->
```python
## Import numpy and pandas

# YOUR CODE HERE
import numpy as np
import pandas as pd
```

<!-- grade_id: cell-ab515cd97836235b, points: 1, grade: true, solution: false -->
```python

```

The two primary data structures of pandas are 
- Series (1-dimensional) and 
- DataFrame (2-dimensional)

Read this [10 min intro to get started](https://pandas.pydata.org/pandas-docs/stable/10min.html)

### Data Loading

You are given two csv files.  
- 'sales1.csv'
- 'sales2.csv'

Load both CSV as dataframes and Investigate the data and its types

<!-- grade_id: cell-15bfe6e4718647be, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
import pandas as pd
import numpy as np
file1=pd.read_csv('sales1.csv')
file2=pd.read_csv('sales2.csv')
```

<!-- grade_id: cell-3c85ac1eab5f8b19, points: 0, grade: false, solution: true -->
```python
# Investigate the data 
# YOUR CODE HERE
print(type(file1))

print(type(file2))
```

### Merge Dataframes
Merge the two dataframes into a single dataframe "sales" based on order id such that new data consist of common order ID and all the other columns.
Also analyse the data in different columns.
Use [Merge function](https://pandas.pydata.org/pandas-docs/version/0.23.4/generated/pandas.merge.html)

<!-- grade_id: cell-82ee3339bf5f4570, points: 0, grade: false, solution: true -->
```python
sales = None
# YOUR CODE HERE

sales=pd.merge(file1,file2,on='Order ID',how='outer')
#print(sales)


```

<!-- grade_id: cell-52d67c318540c515, points: 1, grade: true, solution: false -->
```python
assert(sales[sales["Order ID"]==686800706]["Country"].values == "Libya")
assert(sales[sales["Order ID"]==686800706]["Item Type"].values == "Cosmetics")
```

### Data Cleaning
If you look at the data we can see some issues,
- Some of the Item Type start with " * " we want to remove the asterik in the column
- Some columns data are missing indicated by NaN (u can try isnull(), isna() function)
  - If Region or Country data is missing replace it with unknown
  - if other data are missing, remove it from "sales" dataframe and add it in a "missing" dataframe 

Reference for this section:
- [Working with missing data](https://pandas.pydata.org/pandas-docs/stable/missing_data.html)
- [Python Lambda Expressions](https://docs.python.org/3/tutorial/controlflow.html#lambda-expressions)
- [Pandas Apply Function](https://pandas.pydata.org/pandas-docs/version/0.18/generated/pandas.Series.apply.html)
- [Pandas FillNa](https://pandas.pydata.org/pandas-docs/version/0.17.0/generated/pandas.DataFrame.fillna.html)

<!-- grade_id: cell-36c46b4d53bf011c, points: 0, grade: false, solution: true -->
```python
# Investigate dataframe for Null (NA/NAN) and count the number of rows with NA/NAN

# YOUR CODE HERE

NA_count = sales[sales.isnull().any(axis = 1)].shape[0]
# YOUR CODE HERE


print("Total=" , NA_count)

```

Output:
```
Total= 31
```

<!-- grade_id: cell-c3f7ee60f7aec5bf, points: 1, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-04b3a523548b2603, points: 0, grade: false, solution: true -->
```python
# remove "*", NaN, add missing sales data to "missing" dataframe 
# YOUR CODE HERE

sales['Item Type'] = sales['Item Type'] .apply(lambda x : str(x).replace("*",""))

sales['Region'] = sales['Region'] .fillna("unknown")
sales['Country'] = sales['Country'] .fillna("unknown")


missing = sales[sales.isna().any(axis = 1)]


sales = sales[sales.notnull().all(axis = 1)]
```

<!-- grade_id: cell-7bc1b916b4472109, points: 1, grade: true, solution: false -->
```python
assert(sales[sales["Order ID"]==267066323]["Item Type"].values == "Cereal")
assert(sales[sales["Order ID"]==267066323]["Region"].values == "unknown")
```

Find the number of unique items available.

<!-- grade_id: cell-588bcfca26c28277, points: 0, grade: false, solution: true -->
```python
NumItems = 0
# YOUR CODE HERE
NumItems=len(set(sales))
print(NumItems)
```

Output:
```
12
```

<!-- grade_id: cell-51b5cda4685cd3d0, points: 1, grade: true, solution: false -->
```python

```

```python
sales.groupby('Item Type', as_index = False)['Unit Cost'].describe()
```

### Grouping and Analysis 
Find out the statistical summary for ‘Unit Cost’ column for each Items.
Then extract the count, mean and std for "Cereal","Fruits","Vegetables","Baby Food" and "Meat" as a dataframe "Kitchen"
References:
- [Pandas Indexing](https://pandas.pydata.org/pandas-docs/stable/indexing.html)
- [Groupby](https://pandas.pydata.org/pandas-docs/stable/groupby.html)
- [Loc](https://pandas.pydata.org/pandas-docs/version/0.23.4/generated/pandas.DataFrame.loc.html)
- [filter](https://pandas.pydata.org/pandas-docs/stable/generated/pandas.DataFrame.filter.html)

<!-- grade_id: cell-d5b95b853fd96b60, points: 0, grade: false, solution: true -->
```python
coll = sales.groupby(['Item Type'])['Unit Cost'].describe()
# YOUR CODE HERE
Kitchen=coll.loc[["Cereal","Fruits","Vegetables","Baby Food","Meat"] , [ 'count', ' mean' , 'std']]

```

Output:
```
/opt/conda/lib/python3.6/site-packages/ipykernel_launcher.py:3: FutureWarning: 
Passing list-likes to .loc or [] with any missing label will raise
KeyError in the future, you can use .reindex() as an alternative.

See the documentation here:
https://pandas.pydata.org/pandas-docs/stable/indexing.html#deprecate-loc-reindex-listlike
  This is separate from the ipykernel package so we can avoid doing imports until
```

```python
#
```

<!-- grade_id: cell-986e659a273a364c, points: 1, grade: true, solution: false -->
```python
assert(Kitchen.loc["Cereal"]["count"]==77)
```

<!-- grade_id: cell-346f31cb4ea47c75, points: 0, grade: false, solution: true -->
```python
#Generate the new column ‘Total Profit’ for each orders (Total Profit = Total Revenue - Total Cost)
# YOUR CODE HERE
import pandas as pd
# YOUR CODE HERE
sales['Total Profit'] = sales['Total Revenue'] - sales['Total Cost']
```

<!-- grade_id: cell-6419595bc71f6a09, points: 1, grade: true, solution: false -->
```python
assert(sales[sales["Order ID"]==246222341]["Total Profit"].values == 145419.62)
```

<!-- grade_id: cell-aae9d532c4dabf17, points: 0, grade: false, solution: true -->
```python
# How many data sets are there with profit greater than 60%

import pandas as pd
# YOUR CODE HERE
#ans = None
sales['Total Profit%'] = sales['Total Profit'] /sales['Total Cost'] *100

ans = (sales['Total Profit%']  >= 60).count()

ans




```

Output:
```
972
```

<!-- grade_id: cell-c053463e75781632, points: 1, grade: true, solution: false -->
```python

```

```python

```